---

## <font color='blue'> Scenario: Bob wants to encrypt 16 bytes with AES and send to Alice. First, Bob generates a random key as follows: </font>

```
from Crypto.Cipher import AES
from Crypto.Random import get_random_bytes
AES_key = 8*get_random_bytes(2)
```
## <font color='blue'> Then Bob encrypts two messages (ptxt1 and ptxt2) using the generated AES_key and sends to Alice using the following code: </font>

```
cipher = AES.new(AES_key, AES.MODE_ECB)
ctxt1 = cipher.encrypt(ptxt1).hex()
ctxt2 = cipher.encrypt(ptxt2).hex()
```
## <font color='blue'> However, Eve is monitoring Bob's communications and captures the two ciphertexts (ctxt1 and ctxt2): </font>

```
ctxt1 =  3012e00127813d5a563f775e0906d081
ctxt2 =  79789673ec9cf766126751c0d7486604
```
## <font color='blue'>  program that recovers ptxt1 and ptxt2? </font>

```python

from Cryptodome.Cipher import AES
from Cryptodome.Random import get_random_bytes
import string
```

```python

from Crypto.Cipher import AES
from Crypto.Random import get_random_bytes
import string
```

<!-- grade_id: cell-ee77a78ea2f71eb5, points: 0, grade: false, solution: true -->
```python

def generate_key(seed: int) -> bytes:
   
    key = b""
    key = (seed.to_bytes(2,byteorder='big') * 8) #converting seed to 16 byte
    #raise NotImplementedError()
    return key

def AES_decrypt(key: bytes, ctxt: bytes) -> str:
   
    ptxt = ""
    cipher=AES.new(key, AES.MODE_ECB) # enrcryption
    ptxt=cipher.decrypt(ctxt)  # decryption
    return ptxt.decode('iso-8859-1')


def printable(ptxt: str) -> bool:
   
    answer = False
    l = string.printable
    answer=all(i in l for i in ptxt) # checking each index of ptxt with l
    #raise NotImplementedError()

    return answer


def get_plaintexts(ctxt1: str, ctxt2: str) -> tuple:
    ptxt1 = ""
    ptxt2 = ""

    ctxt1_bytes=bytes.fromhex(ctxt1) # converting ctxt1 to bytes
    ctxt2_bytes=bytes.fromhex(ctxt2) # converting ctxt2 to bytes
    for seed in range(65536):
        key=generate_key(seed)#generate key
        try:
            ptxt1=AES_decrypt(key,ctxt1_bytes) # decrypt ctxt1
            ptxt2=AES_encrypt(key,ctxt2_bytes) # decrypt ctxt2
            if printable(ptxt1) and printable(ptxt2): #if containing all english characters print
                    return (ptxt1, ptxt2)
        except Exception: # handling None
            continue
```

<!-- grade_id: cell-5160b531c3f2ba23, points: 5, grade: true, solution: false -->
```python
key1 = generate_key(432)
key2 = generate_key(2**16-1)
assert key1 == b'\x01\xb0\x01\xb0\x01\xb0\x01\xb0\x01\xb0\x01\xb0\x01\xb0\x01\xb0'
assert key2 == b'\xff\xff\xff\xff\xff\xff\xff\xff\xff\xff\xff\xff\xff\xff\xff\xff'
```

<!-- grade_id: cell-d21d3f20f052308e, points: 5, grade: true, solution: false -->
```python
ctxt1 = "3012e00127813d5a563f775e0906d081"
ctxt2 = "79789673ec9cf766126751c0d7486604"

ptxt1 = AES_decrypt(key1,bytes.fromhex(ctxt1))
ptxt2 = AES_decrypt(key1,bytes.fromhex(ctxt2))

hex1 = ptxt1.encode('iso-8859-1').hex()
hex2 = ptxt2.encode('iso-8859-1').hex()
print(hex1,hex2)
assert hex1 == "f8a29e8998663f024297d30cae922d0c"
assert hex2 == "f033055f92b698c2b715ed64bdc5a909"
```

Output:
```
f8a29e8998663f024297d30cae922d0c f033055f92b698c2b715ed64bdc5a909
```

<!-- grade_id: cell-9a9e68ee3dc794c7, points: 5, grade: true, solution: false -->
```python
ans1 = printable("Isthis,printable!!")
ans2 = printable(bytes.fromhex("f033055f92b698c2b715ed64bdc5a9").decode('iso-8859-1'))
assert ans1 == True
assert ans2 == False
```

<!-- grade_id: cell-72df4c8bb32f11eb, points: 5, grade: true, solution: false -->
```python
ctxt1 = "3012e00127813d5a563f775e0906d081"
ctxt2 = "79789673ec9cf766126751c0d7486604"

ptxt1, ptxt2 = get_plaintexts(ctxt1, ctxt2)
print(ptxt1,ptxt2)
```

Output:
```
my key is weak!! Eve cant read it
```

```python

```