**Show all your work for full credit. Each source code you submit should include detailed comments and instructions on how to run it in order to confirm that it works as expected. If the program that does not run or throws runtime errors, it cannot be graded. You can refer to the programming guidelines from the TAs here: https://tinyurl.com/CPEG-472-672-Programming-Guide/**

**This is an individual assignment and each student should work on their own. Ensure you don't share any code online or with others (note, using Replit, GitHub and similar online platforms can make your code accessible to others).**

**To submit the assignment, you need to use Jupyter Notebook with the provided cell blocks and follow the naming conventions and instructions posted here: https://tinyurl.com/CPEG-472-672-Programming-Guide/**

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel $\rightarrow$ Restart) and then **run all cells** (in the menubar, select Cell $\rightarrow$ Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and section below:

```python
NAME = "Shruthilaya Arun"
#SECTION = "472"
SECTION = "672"
```

---

#### <font color='blue'> Q1 [40 points total - answer all parts] </font>

### <font color='blue'>  Q1 [40 points] Implement a CBC-MAC construction based on AES-128 as illustrated in the lecture and the book. Select two messages M1 and M2 and use your program to compute the corresponding tags T1 and T2. The size of M1 and M2 should be a multiple of the block size of the block cipher. Then, given (M1,T1) and (M2,T2), compute a valid forgery (M3,T3) for your CBC-MAC.

 

```python
import string
from Crypto.Cipher import AES

BLOCK_SIZE = 16
```

```python
import string
from Cryptodome.Cipher import AES

BLOCK_SIZE = 16
```

```python
# you can use this function to xor two bytes, or two strings.
def xor(x,y):
    if type(x) == type(y) == str:
        return "".join([chr(ord(a) ^ ord(b)) for a,b in zip(x,y)])
    elif type(x) == type(y) == bytes:
        return "".join([chr(a ^ b) for a,b in zip(x,y)]).encode('iso-8859-1')
    else:
        print("Error: type mismatch")
```

<!-- grade_id: cell-80e74608889f4fcd, points: 0, grade: false, solution: true -->
```python
# 20 points
def cbc_mac_aes_128(key: bytes, data: bytes) -> str:
    """ Write a CBC-MAC using AES-128 with iv=0 """
    """ Return value as a hex string """
    assert len(key) == BLOCK_SIZE
    assert len(data) % BLOCK_SIZE == 0
    iv = b'\0' * BLOCK_SIZE
    macval = "" # mac as a hex string
    # YOUR CODE HERE
    iv=b'\0'*BLOCK_SIZE # iv
    cipher=AES.new(key,AES.MODE_ECB) 
    mac=iv # set initial value to iv
    for i in range(0,len(data),BLOCK_SIZE):
        block=data[i:i+BLOCK_SIZE] # get block
        mac=cipher.encrypt(xor(mac,block)) # xor block with mac
    #raise NotImplementedError()
    return mac.hex()
```

<!-- grade_id: cell-dc3a472aeb7df34d, points: 5, grade: true, solution: false -->
```python
secret = b"This is a secret"
m1 =     b"This is a CBCMAC"
tag1 = cbc_mac_aes_128(secret, m1)
assert all(c in string.hexdigits for c in tag1)  ##is hex string
assert len(tag1) == 32 # Is 16 bytes (32 hex digits) long
```

<!-- grade_id: cell-aa7bece7696691c5, points: 15, grade: true, solution: false -->
```python
secret = b"This is a secret"
m1 =     b"This is a CBCMAC"
tag1 = cbc_mac_aes_128(secret, m1)
assert tag1 == "03c30100ee85d03d467362724b2ec83f"

m2 = b"This is Second Msg blah blah ..."
tag2 = cbc_mac_aes_128(secret, m2)
assert tag2 == "49f35e02ce4f6147d6260b17cc163a7f"
```

<!-- grade_id: cell-31ef8cde4f2bb1d2, points: 0, grade: false, solution: true -->
```python
# 20 points
def forgery(M1:bytes,T1:str, M2:bytes, T2:str) -> (bytes, str):
    """ Use this function to forge a message. """
    """ You cannot use the secret variable or cmac_aes_128 """
    M3 = b""
    T3 = ""
    # YOUR CODE HERE
    T1_bytes=bytes.fromhex(T1) # hex to bytes
    M2_xor_T1=xor(M2[:BLOCK_SIZE],T1_bytes) #xor m2 first block with t1
    M3=M1+M2_xor_T1+M2[BLOCK_SIZE:]# add m1 to modified block
    T3=T2

    #raise NotImplementedError()
    return (M3,T3) 

```

<!-- grade_id: cell-cd47222fb100e20e, points: 20, grade: true, solution: false -->
```python
secret = b"This is a secret"
m1 = b"This is a CBCMAC"
m2 = b"This is Second Msg blah blah ..."
tag1 = cbc_mac_aes_128(secret, m1)
tag2 = cbc_mac_aes_128(secret, m2)
(m3, tag3) = forgery(m1, tag1, m2, tag2)
assert cbc_mac_aes_128(secret, m3) == tag3
```

```python

```